<img src="images/banner.png" style="width: 100%;">

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from numpy.testing import (assert_equal, assert_almost_equal, 
                           assert_array_equal, assert_allclose, assert_raises)
%matplotlib inline
```

## Problem 1
Create a function `plot_pop` that accepts a list of (case-insensitive) municipality names, reads `Municipality Data - PSA.csv` and returns a `matplotlib` `Figure` that replicates the figure below. It is composed of superimposed line graphs with the $x$-axis being the census date and $y$-axis corresponding to the population. No need to convert the census dates into datetime.

![](pop.png)

<!-- grade_id: cell-9efbd9ce5211857e, points: 0, grade: false, solution: true -->
```python
def plot_pop(municipalities):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c5bdbb508117c39a, points: 1, grade: true, solution: false -->
```python
ax = plot_pop(['CiTy Of MaNiLa', 'City of Makati', 'QUEZON CITY',
                'city of isabela', 'Angeles City'])
ax.figure.canvas.draw()
assert_equal(
    [t.get_text() for t in ax.get_xticklabels()], 
    ['May-10', 'Feb-60', 'May-75', 'May-90', 'May-00', 'May-10', '']
)
assert_equal(
    [t.get_text() for t in ax.get_yticklabels()], 
    ['−0.5', '0.0', '0.5', '1.0', '1.5', '2.0', '2.5', '3.0', '3.5']
)
assert_equal(
    [t.get_text() for t in ax.legend_.texts],
    ['CiTy Of MaNiLa',
     'City of Makati',
     'QUEZON CITY',
     'city of isabela',
     'Angeles City']
)
```

## Problem 2
Create a function `find_max` that accepts a `province` (case-insensitive) and reads `Municipality Data - PSA.csv` then returns the municipality and census dates where the change in population is greatest. Raise `ValueError` if `province` is not found.

Notes:
* In the Philippines, administrative units typically go from Country to Region to Province to City/Municipality to Barangay, although there are exceptions
* Rows with `is_total` = 1 correspond to the total population (excluding the cities) as computed by the Philippines Statistics Authority (PSA). Do not forget to remove these rows when calculating provincial statistics.

<!-- grade_id: cell-816da45ab9261b20, points: 0, grade: false, solution: true -->
```python
def find_max(province):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e33d1d44f3b93e97, points: 1, grade: true, solution: false -->
```python
assert_raises(ValueError, find_max, 'nowhere')
assert_equal(
    find_max('metropolitan manila'), 
    ('QUEZON CITY', 'May-00', 'Aug-07'))
assert_equal(
    find_max('Pampanga'), 
    ('ANGELES CITY', 'May-10', 'Aug-15'))
```

## Problem 3

Create a function `most_populous` that reads `Municipality Data - PSA.csv` and returns a `pandas` `Series` of the 10 provinces with the most mean population for `Aug-15`. The index should be the name of the province and the value should be the mean population. Sort them by decreasing mean population.

<!-- grade_id: cell-c20c8567f112435b, points: 0, grade: false, solution: true -->
```python
def most_populous():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0c0835cb8e11819d, points: 1, grade: true, solution: false -->
```python
mp = most_populous()
assert_array_equal(
    mp[:5].index,
    ['Metropolitan Manila', 'Rizal', 'Davao del Sur', 'Cavite', 'Bulacan']
)
assert_allclose(
    mp[:5].tolist(),
    [757485.4705882353, 206016.2142857143, 161370.0625, 159926.13043478262, 
     137169.625]
)
```

## Problem 4

Create a function `plot_sales_by_country` that reads `/mnt/data/public/retaildata/Online Retail.csv` and returns a matplotlib `Axes` corresponding to the figure below. Sales per transaction is the production of the unit price and quantity.

![](sales.png)

<!-- grade_id: cell-e80dccf6b01c4272, points: 0, grade: false, solution: true -->
```python
def plot_sales_by_country():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6a561931fcae8658, points: 1, grade: true, solution: false -->
```python
ax = plot_sales_by_country()
assert_equal(ax.get_xlabel(), 'Sales (GBP)')
assert_equal(ax.get_ylabel(), 'Country')
assert_equal(
    [text.get_text() for text in ax.get_xticklabels()],
    ['0', '250', '500', '750', '1000', '1250', '1500', '1750', '2000']
)
assert_equal(
    [text.get_text() for text in ax.get_yticklabels()],
    ['Saudi Arabia',
     'Bahrain',
     'Czech Republic',
     'RSA',
     'Brazil',
     'European Community',
     'Lithuania',
     'Lebanon',
     'USA',
     'United Arab Emirates']
)
```

## Problem 5

Create a function `get_invoices_per_month` that reads `/mnt/data/public/retaildata/Online Retail.csv` and returns a `pandas` `Series` where the index is the invoice month and the value is the number of invoices for that month.

<!-- grade_id: cell-ad78e6e872b71bb7, points: 0, grade: false, solution: true -->
```python
def get_invoices_per_month():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-edd08e3163f5eb59, points: 1, grade: true, solution: false -->
```python
invoices = get_invoices_per_month()
assert_equal(len(invoices), 12)
assert_equal(invoices.index.tolist(), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12])
assert_equal(invoices[:5].tolist(), [1914, 1647, 2043, 1821, 2179])
```

## Problem 6

Create a function `get_invoices_by_day_hour` that reads `/mnt/data/public/retaildata/Online Retail.csv` and returns a `pandas` `DataFrame` with index equal to the invoice day of the week, columns equal to invoice hour of day, and values equal to the number of invoices for that day and hour.

<!-- grade_id: cell-56f7140daf0dca28, points: 0, grade: false, solution: true -->
```python
def get_invoices_by_day_hour():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-029d718ead0e793d, points: 1, grade: true, solution: false -->
```python
df_pivot = get_invoices_by_day_hour()
assert_equal(df_pivot.shape, (7, 15))
assert_equal(
    df_pivot.index.tolist(),
    ['Mon', 'Tue', 'Wed', 'Thu', 'Fri', 'Sat', 'Sun']
)
assert_equal(
    df_pivot.columns.tolist(),
    [6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20]
)
assert_equal(
    df_pivot.iloc[0].tolist(),
    [0, 7, 101, 303, 453, 466, 625, 560, 450, 472, 302, 148, 20, 23, 2]
)
```

## Problem 7

Create a function `get_mean_rating_per_year` that reads the first 10000 data lines of `/mnt/data/public/book-crossing/BX-Books.csv` and all of `/mnt/data/public/book-crossing/BX-Book-Ratings.csv` then returns a `pandas` `Series` where the index corresponds to `Year-Of-Publication`, in ascending numerical order, and the values correspond to the mean `Book-Rating` for that year. Both files follow `latin1` encoding. Skip bad lines.

<!-- grade_id: cell-797d2716eaa7034d, points: 0, grade: false, solution: true -->
```python
def get_mean_rating_per_year():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b6fcc4800be404f3, points: 1, grade: true, solution: false -->
```python
mean_ratings = get_mean_rating_per_year()
assert_equal(len(mean_ratings), 62)
assert_equal(
    mean_ratings.index[:10].tolist(),
    [0, 1920, 1937, 1940, 1941, 1942, 1950, 1951, 1952, 1953]
)
assert_allclose(
    mean_ratings[:10].tolist(),
    [3.48782373405489,
     3.96,
     1.6666666666666667,
     2.5,
     7.0,
     2.8125,
     3.8,
     0.0,
     5.5,
     2.5]
)
```

## Problem 8

Create a function `hourly_hashtag` that reads the first 1M data lines of `/mnt/data/public/nowplaying-rs/nowplaying_rs_dataset/user_track_hashtag_timestamp.csv` and returns a pandas data frame with columns `hashtag`, `created_at` and `count`. The input `created_at` column is in UTC. The `count` column contains the number of `created_at`, binned hourly and in Philippine time, for that `hashtag`. Sort by `hashtag` then by `created_at`. 

<!-- grade_id: cell-7db5228e1ccea4eb, points: 0, grade: false, solution: true -->
```python
def hourly_hashtag():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73457b59e6f6796e, points: 1, grade: true, solution: false -->
```python
df_hashtag = hourly_hashtag()
assert_equal(df_hashtag.shape, (73050, 3))
assert_equal(df_hashtag.columns.tolist(), ['hashtag', 'created_at', 'count'])
assert_equal(
    df_hashtag[:10].values.tolist(),
    [['089DJ', pd.Timestamp('2014-01-05 08:00:00+0800', tz='Asia/Manila'), 1],
     ['089DJ', pd.Timestamp('2014-01-05 09:00:00+0800', tz='Asia/Manila'), 2],
     ['089DJ', pd.Timestamp('2014-01-12 08:00:00+0800', tz='Asia/Manila'), 3],
     ['089DJ', pd.Timestamp('2014-01-12 09:00:00+0800', tz='Asia/Manila'), 4],
     ['089DJ', pd.Timestamp('2014-01-12 10:00:00+0800', tz='Asia/Manila'), 9],
     ['089DJ', pd.Timestamp('2014-01-12 11:00:00+0800', tz='Asia/Manila'), 4],
     ['0_Count', pd.Timestamp('2014-01-07 18:00:00+0800', tz='Asia/Manila'), 
      1],
     ['1038FM', pd.Timestamp('2014-01-05 22:00:00+0800', tz='Asia/Manila'), 
      1],
     ['1093KFKU', pd.Timestamp('2014-01-01 14:00:00+0800', tz='Asia/Manila'),
      2],
     ['1093KFKU', pd.Timestamp('2014-01-01 19:00:00+0800', tz='Asia/Manila'), 
      1]]
)
```

## Problem 9

Create a function `aisle_counts` that reads the first 1M data lines of `/mnt/data/public/instacart/instacart_2017_05_01/order_products__prior.csv` and returns a `pandas` `Series` of the number of orders per aisle sorted by decreasing number of orders. Hint: you should look at other files in the directory that you can use.

<!-- grade_id: cell-140254da5863eaba, points: 0, grade: false, solution: true -->
```python
def aisle_counts():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a4f213c69e35ce72, points: 1, grade: true, solution: false -->
```python
ac = aisle_counts()
assert_array_equal(
    ac.index[:10],
    ['fresh fruits', 'fresh vegetables', 'packaged vegetables fruits',
     'yogurt', 'packaged cheese', 'milk', 'water seltzer sparkling water',
     'chips pretzels', 'soy lactosefree', 'bread']
)
assert_array_equal(
    ac[:10].to_list(),
    [112409, 105047, 54781, 44674, 30184, 27428, 25789, 22190, 19839, 18091]
)
```

## Problem 10

Create a function `from_to` that reads the first 1000 data lines of `/mnt/data/public/wikipedia/clickstream/clickstream/2017-11/clickstream-enwiki-2017-11.tsv.gz` and returns a `pandas` `DataFrame` where the index are the unique values of the source column (first column) sorted in lexicographical order, the columns are the unique values of the destination column (second column) sorted in lexicographical order and the values are the corresponding views (fourth column). Set the value to 0 if it wasn't viewed.

<!-- grade_id: cell-48038f4318286406, points: 0, grade: false, solution: true -->
```python
def from_to():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7458d465b09ac705, points: 1, grade: true, solution: false -->
```python
df_ft = from_to()
assert_equal(df_ft.shape, (621, 171))
assert_array_equal(
    df_ft.index[:10],
    ['1._FSV_Mainz_05', '13_Reasons_Why', '1939_in_film', '1970_in_music',
     '1976_Formula_One_season', '1976_in_Irish_television',
     '1977_in_Ireland', '1977_in_television', '1980s', '1990_in_film']
)
assert_array_equal(
    df_ft.columns[:10],
    ['1921_Navy_Midshipmen_football_team', '1927_in_Turkey',
     '1947–48_Scottish_Division_Three', '1958_in_association_football',
     '1977_in_Irish_television',
     "2002_Pan_Pacific_Swimming_Championships_–_Women's_50_metre_freestyle",
     '2010_Sony_Ericsson_Open', '2011_UCI_World_Tour',
     '2012_Oregon_State_Beavers_football_team',
     '2015_World_Polo_Championship']
)
```

<img src="images/banner-down.png" style="width: 100%;">