<img src="images/banner.png" style="width: 100%;">

```python
import numpy as np
import pandas as pd
from numpy.testing import assert_equal, assert_allclose
```

## Problem 1

Create a function `double_work` that accepts a pandas data frame `df` with a time series index and columns `a` and `b`. It modifies `df` such that the value of `a` from 9am to 5pm (inclusive) of 1 Jan 2021 is double the value of `b` for those hours.

<!-- grade_id: cell-0acb5a75f001754f, points: 0, grade: false, solution: true -->
```python
def double_work(df):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-17550f5d3c0e0259, points: 1, grade: true, solution: false -->
```python
df = pd.DataFrame(
    {
        'a': range(24), 
        'b': np.linspace(1, 100, 24)
    },
        index=pd.date_range('2021-01-01 00:00', '2021-01-01 23:00', freq='H')
)
double_work(df)
assert_equal(
    df.index[:10].tolist(),
    [pd.Timestamp('2021-01-01 00:00:00'),
     pd.Timestamp('2021-01-01 01:00:00'),
     pd.Timestamp('2021-01-01 02:00:00'),
     pd.Timestamp('2021-01-01 03:00:00'),
     pd.Timestamp('2021-01-01 04:00:00'),
     pd.Timestamp('2021-01-01 05:00:00'),
     pd.Timestamp('2021-01-01 06:00:00'),
     pd.Timestamp('2021-01-01 07:00:00'),
     pd.Timestamp('2021-01-01 08:00:00'),
     pd.Timestamp('2021-01-01 09:00:00')]
)
assert_allclose(
    df.iloc[:10].values,
    [[ 0.        ,  1.        ],
     [ 1.        ,  5.30434783],
     [ 2.        ,  9.60869565],
     [ 3.        , 13.91304348],
     [ 4.        , 18.2173913 ],
     [ 5.        , 22.52173913],
     [ 6.        , 26.82608696],
     [ 7.        , 31.13043478],
     [ 8.        , 35.43478261],
     [79.47826087, 39.73913043]]
)
```

## Problem 2

Create a function `hourly_mem_usage` that reads `mem.csv` and returns the mean hourly memory usage in Philippine time (UTC+8) as a pandas time series. The first column of `mem.csv` is the time in UTC and the second column is the memory usage in bytes. Each interval is closed on the left but open on the right.

<!-- grade_id: cell-2ab74f3afc7e2594, points: 0, grade: false, solution: true -->
```python
def hourly_mem_usage():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2df0e3464012ace8, points: 1, grade: true, solution: false -->
```python
hmu = hourly_mem_usage()
assert_equal(len(hmu), 49)
assert_equal(
    hmu.index[:10].tolist(),
    [pd.Timestamp('2021-05-28 00:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 01:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 02:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 03:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 04:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 05:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 06:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 07:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 08:00:00+0800', tz='Asia/Manila'),
     pd.Timestamp('2021-05-28 09:00:00+0800', tz='Asia/Manila')]
)
assert_allclose(
    hmu[:10].tolist(),
    [369325498368.0,
     389165290837.3333,
     394183460454.4,
     388326653132.8,
     390690503065.6,
     391548573832.5333,
     389338755072.0,
     389219362406.4,
     390718947328.0,
     375491907037.86664]
)
```

## Problem 3

Create a function `daily_mem_usage` that reads `mem.csv` and returns the mean daily memory usage as a pandas time series with a daily period index. The first column of `mem.csv` is the time and the second column is the memory usage in bytes. Each interval is closed on the left but open on the right.

<!-- grade_id: cell-702732e11f1d52be, points: 0, grade: false, solution: true -->
```python
def daily_mem_usage():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-875a1de1e6697679, points: 1, grade: true, solution: false -->
```python
dmu = daily_mem_usage()
assert_equal(
    dmu.index[:2].tolist(),
    [pd.Period('2021-05-27', 'D'), pd.Period('2021-05-28', 'D')]
)
assert_allclose(
    dmu[:2].tolist(),
    [390154853588.5283, 346686658616.8889]
)
```

## Problem 4

Create a function `longest_distances` that reads the first 1M data lines of `/mnt/data/public/nyctaxi/trip_data/trip_data_1.csv` and returns a `pandas.Series` containing the maximum `trip_distance` for each pickup hour of day and passenger count. The index is a hierarchical index of `pickup_datetime` and `passenger_count`.

<!-- grade_id: cell-5f238bcbd72181a3, points: 0, grade: false, solution: true -->
```python
def longest_distances():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-edd54e2dfb967fd5, points: 1, grade: true, solution: false -->
```python
distances = longest_distances()
assert_equal(isinstance(distances, pd.core.series.Series), True)
assert_equal(distances.shape, (149,))
assert_equal(distances.index.levshape, (24, 7))
assert_equal(distances.index.names, ['pickup_datetime', 'passenger_count'])
```

## Problem 5
Create a function `mean_ratings` that reads `/mnt/data/public/insideairbnb/data.insideairbnb.com/united-kingdom/england/london/2015-04-06/data/listings.csv.gz` and returns a `pandas.Series` containing the average `review_scores_ratings` of hosts, binned by monthly `host_since`. Ignore `NA`s or null values when computing the mean.

<!-- grade_id: cell-990d37718690cffb, points: 0, grade: false, solution: true -->
```python
def mean_ratings():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1c32dec1fa166f06, points: 1, grade: true, solution: false -->
```python
import datetime
ratings = mean_ratings()
assert_equal(isinstance(ratings, pd.core.series.Series), True)
assert_equal(ratings.shape, (80,))
assert_equal(ratings.index[0], datetime.datetime(2008, 9, 30))
assert_equal(ratings.index[-1], datetime.datetime(2015, 4, 30))
```

## Problem 6

Create a function `product_aisles` that reads `/mnt/data/public/instacart/instacart_2017_05_01/products.csv` and returns a `pandas.DataFrame`. The data frame should have the contents of that file but with an additional column `aisle` that corresponds to the name of the aisle according to `/mnt/data/public/instacart/instacart_2017_05_01/aisles.csv`. The index should be the `product_id` column.

<!-- grade_id: cell-40c2bc95fd97b5af, points: 0, grade: false, solution: true -->
```python
def product_aisles():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4e4a5307f3124e7c, points: 1, grade: true, solution: false -->
```python
df_products_aisles = product_aisles()
assert_equal(isinstance(df_products_aisles, pd.DataFrame), True)
assert_equal(
    df_products_aisles.columns.tolist(), 
    ['product_name', 'aisle_id', 'department_id', 'aisle'])
assert_equal(df_products_aisles.shape, (49688, 4))
```

## Problem 7

Create a function `tracks_with_loc` that reads `/mnt/data/public/millionsong/AdditionalFiles/unique_tracks.txt` and returns a `pandas` `DataFrame` that contains the contents of the file but with the artist location from `/mnt/data/public/millionsong/AdditionalFiles/artist_location.txt` matched based on the artist name. Leave the value as null if the artist is not found in `artist_location.txt`. Use the first entry in the file if an artist has more than one location. Sort the data frame by increasing track ID.

<!-- grade_id: cell-378eb694f1a79212, points: 0, grade: false, solution: true -->
```python
def tracks_with_loc():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-74c5a21523df3512, points: 1, grade: true, solution: false -->
```python
df_tracks_loc = tracks_with_loc()
assert_equal(isinstance(df_tracks_loc, pd.DataFrame), True)
assert_equal(df_tracks_loc.shape, (1000000, 8))
assert_equal(
    df_tracks_loc.columns.tolist(), 
    ['track_id', 'song_id', 'artist', 'title', 
     'artist_id', 'lat', 'lon', 'location'])
```

## Problem 8
Create a function `party_votes` that reads `/mnt/data/public/elections/comelec/congress_results/congressional_results_2013.csv` and returns a `pandas` `DataFrame` where the index is the `province_or_city`, the columns are the `party_affiliation` and the values are the total number of votes as integer.

<!-- grade_id: cell-c5bd8a8963f5b048, points: 0, grade: false, solution: true -->
```python
def party_votes():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c609d570f23233cf, points: 1, grade: true, solution: false -->
```python
df_votes = party_votes()
assert_equal(isinstance(df_votes, pd.DataFrame), True)
assert_equal(df_votes.shape, (108, 24))
assert_equal((df_votes.dtypes == int).all(), True)
assert_equal(df_votes.index[0], 'Abra')
assert_equal(df_votes.columns[0], 'BPP')
assert_equal(df_votes.iloc[0,0], 0)
```

## Problem 9
Create a function `naia_traffic` that reads `/mnt/data/public/opendata/transport/caap-aircraft/airdata_aircraft_movement_2016.csv` and returns the number of passengers per month, as integer, for every `airline_operator` in `NAIA` as a `pandas` `DataFrame`. The months should be in title case.

<!-- grade_id: cell-27d7c6647b2e5fc3, points: 0, grade: false, solution: true -->
```python
def naia_traffic():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c75134a88e986f33, points: 1, grade: true, solution: false -->
```python
df_naia = naia_traffic()
assert_equal(isinstance(df_naia, pd.DataFrame), True)
assert_equal(
    df_naia.columns.tolist(), 
    ['airline_operator', 'month', 'passengers'])
assert_equal(df_naia.shape, (36, 3))
df_naia_list = df_naia.to_numpy().tolist()
assert_equal(['Domestic', 'April', 13517] in df_naia_list, True)
assert_equal(['G. Aviation', 'April', 3586] in df_naia_list, True)
assert_equal(['International', 'April', 8587] in df_naia_list, True)
```

## Problem 10

Create a function `pudo` that reads the first 1M data lines of `/mnt/data/public/nyctaxi/all/yellow_tripdata_2017-12.csv` and returns a `pandas` `DataFrame` where the index is the unique values of `PULocationID`, the columns are the unique values of `DOLocationID`, and the values are the number of times the `PULocationID`-`DOLocationID` pair occurred.

<!-- grade_id: cell-5107e45cf1b18dd9, points: 0, grade: false, solution: true -->
```python
def pudo():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b20289a550a9cad2, points: 1, grade: true, solution: false -->
```python
df_pudo = pudo()
assert_equal(df_pudo.shape, (243, 261))
assert_equal(df_pudo.index[:10].tolist(), [1, 3, 4, 5, 6, 7, 8, 9, 10, 11])
nonzeros = df_pudo.to_numpy().nonzero()
assert_equal(nonzeros[0][:10], [0, 0, 0, 1, 1, 1, 1, 1, 1, 2])
assert_equal(nonzeros[1][:10], [0, 128, 260, 2, 31, 68, 133, 155, 164, 0])
assert_equal(
    df_pudo.iloc[0][:10].to_numpy(), 
    [86,  0,  0,  0,  0,  0,  0,  0,  0,  0])
```

<img src="images/banner-down.png" style="width: 100%;">